# Instructions

**Please follow these instructions carefully to ensure your submission is graded correctly:**

- Fill in **your answers** in any place that says "YOUR CODE HERE" or "YOUR ANSWER HERE" in the provided `.ipynb` files.

- Your answers to code-based questions must be **generated from code**. Hardcoded answers will result in a deduction of points.  

- Do not change the **autograding cells** labelled with `"AUTOGRADING CELL - DO NOT TOUCH!"`. Modifying or deleting them may lead to incorrect grading results.

- All plots must have appropriate **axis labels** including units where applicable. Failing to do so will result in a partial deduction of points.

- Always proide **units** for quantities where applicable. Failing to do so will result in a partial deduction of points.

- The **code quality** will not be graded. However, concise, well-structured code makes it easier to awared partial points in case of errors.

- The quality of **written answers** and **figures** will be considered during grading.

- Do not use additional packages that are not part of the course **environment**. This would lead to errors during autograding.

- Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel $\rightarrow$ Restart) and then **run all cells** (in the menubar, select Cell $\rightarrow$ Run All).

- Submissions must be your own work, **plagiarism** from the web or peers will be checked for and sanctioned.

- Please upload only the provided `.ipynb` files, filled in with your answers.

- Please **do not** put the `.ipynb` files into archives (e.g., `.zip`, `.tar.gz`); also `.py` files will **not** be accepted. Please **do not change the file names** and do not upload additional files.

- Submission on ISIS *(for TU Berlin students)* requires **two-factor authentication**.

---

## Task 2: Danish Electricity Market Time Series (23 points)

Under the link below (`url`), you will find a CSV file with hourly time series on the electricity demand (in MW), the wind and solar generation (in MW), and the electricity prices (in €/MWh) for the market region DK1 (Western Denmark).

For the tasks listed below, all requested plots must be created with `matplotlib`. For all plots, choose suitable figure sizes, axis labels and line styles. Partial points will be deducted for incomplete or unclear plots.

```python
url = "https://tubcloud.tu-berlin.de/s/qR5YN7XDfRSZNWQ/download/time-series.csv"
```

```python
import numpy as np
import pandas as pd
import matplotlib.pyplot as plt
```

**a)** Read the dataset into a `pandas.DataFrame` named `df`. The first column should parsed into the index as `pandas.DatetimeIndex`. **(2 points)**

<!-- grade_id: cell-2a85756ef9f71252, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: 2-a-t1, points: 1, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
assert "df" in globals()
assert isinstance(df, pd.DataFrame)
```

<!-- grade_id: 2-a-t2, points: 1, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
```

**b)** What are the average values for each column in the dataset? Return the solution as a `pandas.Series`. **(2 points)**

<!-- grade_id: cell-e89931fa09a5a0f3, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: 2-b-t1, points: 2, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
assert isinstance(_, pd.Series)
```

**c)** Normed by the maximum value of each column, which column has the highest standard deviation? **(3 points)**

<!-- grade_id: cell-be0160a3bb835e27, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: 2-c-t1, points: 3, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
assert isinstance(_, str)
```

**d)** Find the timestamps where the electricity price is above 100 €/MWh. Return all hours where this is the case as a `pandas.DatetimeIndex`. **(3 points)**

<!-- grade_id: cell-6c64a98723c08c05, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: 2-d-t1, points: 3, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
assert isinstance(_, pd.DatetimeIndex)
```

**e)** What percentage of hours does the sum of wind and solar generation exceed the electricity demand (in %)? **(3 points)**

<!-- grade_id: cell-a1b1e555c7980a9c, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: 2-e-t1, points: 2, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
assert isinstance(_, float)
```

<!-- grade_id: 2-e-t2, points: 1, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
```

**f)** What is the Pearson correlation coefficient between wind generation and electricity prices? **(2 point)**

<!-- grade_id: cell-971a81b5f7d95a68, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: 2-f-t1, points: 2, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
assert isinstance(_, float)
```

**g)** Plot a scatter plot of wind generation against electricity prices. **(2 points)**

<!-- grade_id: 2-g-t1, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
```

**h)** Plot the hourly time series of wind generation, solar generation and electricity demand for the month of September. **(2 points)**

<!-- grade_id: 2-h-t1, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
```

**i)** Sort the values in descending order and plot the price duration curve. The x-axis should be given as a percentage of the year. **(4 points)**

<!-- grade_id: 2-i-t1, points: 4, grade: true, solution: true -->
```python
# YOUR CODE HERE
```

```python

```