# Instructions

**Please follow these instructions carefully to ensure your submission is graded correctly:**

- Fill in **your answers** in any place that says "YOUR CODE HERE" or "YOUR ANSWER HERE" in the provided `.ipynb` files.

- Your answers to code-based questions must be **generated from code**. Hardcoded answers will result in a deduction of points.  

- Do not change the **autograding cells** labelled with `"AUTOGRADING CELL - DO NOT TOUCH!"`. Modifying or deleting them may lead to incorrect grading results.

- All plots must have appropriate **axis labels** including units where applicable. Failing to do so will result in a partial deduction of points.

- Always proide **units** for quantities where applicable. Failing to do so will result in a partial deduction of points.

- The **code quality** will not be graded. However, concise, well-structured code makes it easier to awared partial points in case of errors.

- The quality of **written answers** and **figures** will be considered during grading.

- Do not use additional packages that are not part of the course **environment**. This would lead to errors during autograding.

- Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel $\rightarrow$ Restart) and then **run all cells** (in the menubar, select Cell $\rightarrow$ Run All).

- Submissions must be your own work, **plagiarism** from the web or peers will be checked for and sanctioned.

- Please upload only the provided `.ipynb` files, filled in with your answers.

- Please **do not** put the `.ipynb` files into archives (e.g., `.zip`, `.tar.gz`); also `.py` files will **not** be accepted. Please **do not change the file names** and do not upload additional files.

- Submission on ISIS *(for TU Berlin students)* requires **two-factor authentication**.

---

## Task 3: Nuclear Power Plants in Global Energy Monitor (55 points)

In this task, we are going to look at the [Global Nuclear Power Tracker](https://globalenergymonitor.org/projects/nuclear-power-tracker/) published by [Global Energy Monitor](https://globalenergymonitor.org/) with a focus on the locations and capacities of nuclear power plants worldwide.

```python
url = "https://tubcloud.tu-berlin.de/s/d5Qt3JWNEseMNGN/download/Global-Nuclear-Power-Tracker-September-2025.xlsx"
```

```python
import cartopy
import numpy as np
import pandas as pd
import geopandas as gpd
import cartopy.crs as ccrs
import matplotlib.pyplot as plt
```

**a)** Read the Excel file from the link above (`url`) with `pandas` and use the provided coordinates to convert it into a `geopandas.GeoDataFrame` named `gdf` with the appropriate coordinate reference system (CRS). Set the index to the "GEM unit ID" column . **(4 points)**

<!-- grade_id: cell-3a904633ad3e0a8f, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: 3-a-t1, points: 1, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
assert isinstance(gdf, gpd.GeoDataFrame)
```

<!-- grade_id: 3-a-t2, points: 1, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
```

<!-- grade_id: 3-a-t3, points: 1, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
```

<!-- grade_id: 3-a-t4, points: 1, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
```

**b)** Find the operating nuclear power plants with a start year of 2025 and list their project names as a Python list or `numpy.ndarray`. **(2 points)**

<!-- grade_id: cell-c3ce939c87be748e, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: 3-b-t1, points: 1, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
assert isinstance(_, list | np.ndarray)
```

<!-- grade_id: 3-b-t2, points: 1, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
```

**c)** Find and list the six countries with the largest announced nuclear power plant capacity (in GW) as a `pandas.Series` with countries as the index and capacities as the values. **(3 points)**

<!-- grade_id: cell-a9d592aa32f8d8b2, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: 3-c-t1, points: 2, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
assert isinstance(_, pd.Series)
```

<!-- grade_id: 3-c-t2, points: 1, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
```

**d)** What share (in %) of the operating nuclear power plant fleet has a determined retirement year? **(3 points)**

<!-- grade_id: cell-6f2520c9c0ca47a4, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: 3-d-t1, points: 1.5, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
assert isinstance(_, float)
```

<!-- grade_id: 3-d-t2, points: 1.5, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
```

**e)** Plot a bar chart which shows the capacity per country (in GW electric) that will be retired by the end of 2045 in ascending order. For power plants with no set retirement year, assume a lifetime of 45 years. For all other power plants, use the provided retirement year. **(5 points)**

<!-- grade_id: 3-e-t1, points: 5, grade: true, solution: true -->
```python
# YOUR CODE HERE
```

**f)** For operating power plants with given information on the construction start date and the commercial operation data, compute the median construction time in years. Use the function `pd.to_datetime` to convert the date columns to datetime objects. **(4 points)**

<!-- grade_id: cell-5cc66464fd8962b5, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: 3-f-t1, points: 2, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
assert isinstance(_, float)
```

<!-- grade_id: 3-f-t2, points: 2, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
```

**g)** Compute the capacity-weighted mean age (in years) of all operating nuclear power plants. **(3 points)**

<!-- grade_id: cell-29dadef51049fd52, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: 3-g-t1, points: 2, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
assert isinstance(_, float)
```

<!-- grade_id: 3-g-t2, points: 1, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
```

**h)** For each country, compute the average age of its operating nuclear power plants weighted by their operational capacity.
Return the mean ages as a `pandas.Series` in ascending order with the country names as the index and the mean ages as the values. **(4 points)**

<!-- grade_id: cell-f3cab5998a2b5f0a, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: 3-h-t1, points: 1, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
assert isinstance(_, pd.Series)
assert "France" in _.index
```

<!-- grade_id: 3-h-t2, points: 1.5, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
```

<!-- grade_id: 3-h-t3, points: 1.5, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
```

**i)** List the names of the operators of the operating Canadian nuclear power plants in descending order of total capacity operated. Return a `pandas.Series` with operator names as the index and total capacities (in GW) as the values. **(3 points)**

<!-- grade_id: cell-05c1439c095f6cef, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: 3-i-t1, points: 2, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
assert isinstance(_, pd.Series)
```

<!-- grade_id: 3-i-t2, points: 0.5, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
```

<!-- grade_id: 3-i-t3, points: 0.5, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
```

**j)** Find the most common operational reactor type (by number of reactors) of each country with operational nuclear power plants. Return a `pandas.Series` with countries as the index and reactor types as the values. **(3 points)**

<!-- grade_id: cell-293ee6e752fcc340, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: 3-j-t1, points: 1, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
assert isinstance(_, pd.Series)
assert "China" in _.index
```

<!-- grade_id: 3-j-t2, points: 2, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
```

**k)** For each country, find the most common operating reactor model by total capacity. Return a `pandas.Series` with countries as the index and the most common reactor model as the values. **(4 points)**

<!-- grade_id: cell-c13b794cd98a7fec, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: 3-k-t1, points: 1, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
assert isinstance(_, pd.Series)
assert "France" in _.index
```

<!-- grade_id: 3-k-t2, points: 0.5, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
```

<!-- grade_id: 3-k-t3, points: 2.5, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
```

**l)** Visualise the locations of operating nuclear power plants on a map covering roughly Pakistan, India, China, Japan, South Korea and North Korea. Ensure the following criteria are met: **(5 points)**

- Use the Mercator projection and a figure size of approximately 15x10 inches.
- The size of the points should be proportional to the capacity of the power plants (no legend required).
- The color of the points should indicate the start year of the power plants.
- The map should include coastlines and country borders; the background should be white.

<!-- grade_id: 3-l-t1, points: 5, grade: true, solution: true -->
```python
# YOUR CODE HERE
```

**In the following, we additionally want to aggregate data on nuclear power plants in Europe to various NUTS levels. The NUTS region data can be read with `geopandas` using the URL (`url_nuts`) below:**

```python
url_nuts = "https://tubcloud.tu-berlin.de/s/5WJ6pGsBKR7a3sp/download/NUTS_RG_10M_2021_4326.geojson"
```

**m)** Calculate the area (in square kilometers) of each NUTS 2 region and return it as a `pandas.Series` with the NUTS 2 region codes as the index (sorted alphabetically) and the areas as the values. **(4 points)**

<!-- grade_id: cell-81ccda725f042f34, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: 3-m-t1, points: 0.5, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
assert isinstance(_, pd.Series)
```

<!-- grade_id: 3-m-t2, points: 0.5, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
```

<!-- grade_id: 3-m-t3, points: 3, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
```

**n)** Find the three NUTS 2 regions with the highest total capacity of operating nuclear power plants. Return a `pandas.Series` with the NUTS 2 region names (`NAME_LATN`, not the `NUTS_ID`) as the index and the total capacities (in MW) as the values. **(3 points)**

<!-- grade_id: cell-ba8338f8666977a4, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: 3-n-t1, points: 2, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
assert isinstance(_, pd.Series)
```

<!-- grade_id: 3-n-t2, points: 1, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
```

**o)** Plot a choropleth map of NUTS3 regions with `matplotlib` which shows each region's total announced nuclear power plant capacity (in MW). The following criteria must be met: **(5 points)**

- The figure size should be approximately 15x10 inches.
- Use the EqualEarth projection.
- Coastlines and country borders should be included in black.
- Set the geographical bounds to focus on continental Europe.
- Use the colormap `"plasma"` starting from value 0.
- Include an appropriately labeled colorbar to the right-hand side of the plot.
- Regions with no announced nuclear power plants should be grey and hatched.

<!-- grade_id: 3-o-t1, points: 5, grade: true, solution: true -->
```python
# YOUR CODE HERE
```

```python

```