# Instructions

**Please follow these instructions carefully to ensure your submission is graded correctly:**

- Fill in **your answers** in any place that says "YOUR CODE HERE" or "YOUR ANSWER HERE" in the provided `.ipynb` files.

- Your answers to code-based questions must be **generated from code**. Hardcoded answers will result in a deduction of points.  

- Do not change the **autograding cells** labelled with `"AUTOGRADING CELL - DO NOT TOUCH!"`. Modifying or deleting them may lead to incorrect grading results.

- All plots must have appropriate **axis labels** including units where applicable. Failing to do so will result in a partial deduction of points.

- Always proide **units** for quantities where applicable. Failing to do so will result in a partial deduction of points.

- The **code quality** will not be graded. However, concise, well-structured code makes it easier to awared partial points in case of errors.

- The quality of **written answers** and **figures** will be considered during grading.

- Do not use additional packages that are not part of the course **environment**. This would lead to errors during autograding.

- Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel $\rightarrow$ Restart) and then **run all cells** (in the menubar, select Cell $\rightarrow$ Run All).

- Submissions must be your own work, **plagiarism** from the web or peers will be checked for and sanctioned.

- Please upload only the provided `.ipynb` files, filled in with your answers.

- Please **do not** put the `.ipynb` files into archives (e.g., `.zip`, `.tar.gz`); also `.py` files will **not** be accepted. Please **do not change the file names** and do not upload additional files.

- Submission on ISIS *(for TU Berlin students)* requires **two-factor authentication**.

---

## Task 1: Location Choice of Offshore Wind Turbine (22 points)

You have been tasked with choosing between two potential locations for a wind
turbine, one in the North Sea and the other in the Baltic Sea. To execute this
task, you have been given a dataset with time series of wind speeds at the two
sites under consideration at ten meters above the ground in units of m/s.
Moreover, your client has already determined the wind turbine model. The power
curve of the turbine is shown in the table below. The turbine has a hub height
of 110 meters and a nominal rating of 4 MW. The cut-out speed beyond which the
wind turbine is shut down is 23 m/s.

**Power curve:**

| Wind Speed [m/s] | Power [MW] |
|------------------:|-----------:|
| < 4              | 0.0        |
| 4–6              | 0.3        |
| 6–8              | 0.9        |
| 8–10             | 2.0        |
| 10–12            | 3.4        |
| 12–16            | 3.9        |
| > 16             | 4          |

**Wind speed time series:**

```python
url = "https://tubcloud.tu-berlin.de/s/PJPjRcqgxTH5wdW/download/wind-speeds.csv"
```

**a)** Import `pandas` under the alias `pd`, `numpy` under the alias `np`, and `matplotlib` like in the tutorials. **(1 point)**

<!-- grade_id: cell-e89931fa09a5a0f3, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: 1-a-t1, points: 1, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
```

**b)** Read the CSV file given under the link (`url` variable) above into a `pandas.DataFrame` such that the timestamps in the first column become the indices and ensure the timestamps are treated as `pd.DatetimeIndex`. Return the DataFrame. **(2 points)**

<!-- grade_id: cell-7ec6376d8f6cef05, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: 1-b-t1, points: 1, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
assert isinstance(_, pd.DataFrame)
```

<!-- grade_id: 1-b-t2, points: 1, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
```

**c)** Convert the wind speeds to a hub height of 110 meters using the log law, assuming open sea as terrain. Return the conversion factor as well as the converted wind speeds of the last ten hours of the year as a `pandas.DataFrame`. **(3 points)**

**Conversion factor:**

<!-- grade_id: cell-6dd2e1a1d080900f, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: 1-c-t1, points: 1.5, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
assert isinstance(_, float)
```

**Converted wind speeds:**

<!-- grade_id: cell-c6f23ab3f8bc3f66, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: 1-c-t2, points: 1.5, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
assert isinstance(_, pd.DataFrame)
assert len(_) == 10
```

**d)** Write a Python function (`power_curve`) that takes the wind speed `v` of one hour as an argument and converts a given single value into power production using the turbine's power curve. **(4 points)**

<!-- grade_id: cell-fde0ac05fede59fc, points: 0, grade: false, solution: true -->
```python
def power_curve(v):
    # YOUR CODE HERE
```

<!-- grade_id: 1-d-t1, points: 2, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
assert power_curve(9) == 2.0
```

<!-- grade_id: 1-d-t2, points: 2, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
```

**e)** Use this function to convert from wind speeds to power production for the full year for each of the two locations under consideration. Return the last ten hours of electricity generation for each location in MW as a `pandas.DataFrame`. **(1 point)**

<!-- grade_id: cell-4d6a79aca87d0f18, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: 1-e-t1, points: 1, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
assert isinstance(_, pd.DataFrame)
assert _.columns.symmetric_difference(["north_sea_m_s", "baltic_sea_m_s"]).empty
assert len(_) == 10
```

**f)** Calculate the total annual energy production in MWh for the North Sea. **(2 points)**

<!-- grade_id: cell-184210ef40ae713d, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: 1-f-t1, points: 2, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
assert isinstance(_, float)
```

**g)** Normalise the power output time series in the North Sea by the turbine's rated capacity and return the last 10 hours as a `pandas.Series`. **(2 points)**

<!-- grade_id: cell-7cd0229ca8cc93a0, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: 1-g-t1, points: 2, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
assert isinstance(_, pd.Series)
assert len(_) == 10
```

**h)** What are the monthly average capacity factors for the North Sea location? Return the capacity factors per month as a `pandas.Series`. **(2 points)**

<!-- grade_id: cell-c38acc9a87e6ae9f, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: 1-h-t1, points: 2, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
assert isinstance(_, pd.Series)
assert len(_) == 12
```

**i)** What share of time is the turbine producing at full capacity in the North Sea? **(2 points)**

<!-- grade_id: cell-1f86d9d797e4f678, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-82665eb64e5db3d1, points: 2, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
assert isinstance(_, int | float)
```

**j)** Plot the capacity factor time series for the full year using `matplotlib`. Choose a figure size of 12 by 3 inches and include appropriate axis labels! **(3 points)**

<!-- grade_id: cell-d0379da2a5d1b939, points: 3, grade: true, solution: true -->
```python
# YOUR CODE HERE
```

```python

```