# Instructions

**Please follow these instructions carefully to ensure your submission is graded correctly:**

- Fill in **your answers** in any place that says "YOUR CODE HERE" or "YOUR ANSWER HERE" in the provided `.ipynb` files.

- Your answers to code-based questions must be **generated from code**. Hardcoded answers will result in a deduction of points.  

- Do not change the **autograding cells** labelled with `"AUTOGRADING CELL - DO NOT TOUCH!"`. Modifying or deleting them may lead to incorrect grading results.

- All plots must have appropriate **axis labels** including units where applicable. Failing to do so will result in a partial deduction of points.

- Always proide **units** for quantities where applicable. Failing to do so will result in a partial deduction of points.

- The **code quality** will not be graded. However, concise, well-structured code makes it easier to awared partial points in case of errors.

- The quality of **written answers** and **figures** will be considered during grading.

- Do not use additional packages that are not part of the course **environment**. This would lead to errors during autograding.

- Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel $\rightarrow$ Restart) and then **run all cells** (in the menubar, select Cell $\rightarrow$ Run All).

- Submissions must be your own work, **plagiarism** from the web or peers will be checked for and sanctioned.

- Please upload only the provided `.ipynb` files, filled in with your answers.

- Please **do not** put the `.ipynb` files into archives (e.g., `.zip`, `.tar.gz`); also `.py` files will **not** be accepted. Please **do not change the file names** and do not upload additional files.

- Submission on ISIS *(for TU Berlin students)* requires **two-factor authentication**.

---

## Task 2: Dimensioning Lossless Storage (41 points)

Consider the time series on wind and solar capacity factors and electricity demand in the Czech
Republic in 2019:

```python
url = "https://tubcloud.tu-berlin.de/s/r6KpCsmwx5wA5A8/download/time-series-cz.csv"
```

The goal of this task is to dimension a lossless storage for different generation mixes of onshore wind and
solar to supply demand. This means that you should assume, for simplicity, that the storing electricity does not incur any losses.

```python
import numpy as np
import pandas as pd
import matplotlib.pyplot as plt
```

**a)** Load the provided dataset as a `pandas.DataFrame` under the name `df` and such that its index contains the time stamps. Return the resulting `pandas.DataFrame`. **(1 point)**

<!-- grade_id: cell-933f67d83f385de5, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: 2-a-t1, points: 1, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
assert isinstance(_, pd.DataFrame)
```

**b)** Convert the unit of the load to GW in a new column `load_gw` and rescale the onshore wind and solar capacity factor time series such that _for each of them_ the annual generation equals the annual load. Store the resulting time series in GW in new columns `wind_gw` and `pv_gw`, respectively. Return the modified `pandas.DataFrame`. **(4 points)**

<!-- grade_id: cell-0befaa7e7547a378, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: 2-b-t1, points: 1, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
assert isinstance(_, pd.DataFrame)
```

<!-- grade_id: 2-b-t2, points: 1.5, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
```

<!-- grade_id: 2-b-t3, points: 1.5, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
```

**c)** Determine the residual load time series for the case where all power is supplied by (1) wind power only and (2) solar power only. 

1. Wind-only case. Return the result as a `pandas.Series` in GW. **(2 points)**

<!-- grade_id: cell-e9939f7537a1ab79, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: 2-c-t1, points: 2, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
assert isinstance(_, pd.Series)
```

2. Solar-only case. Return the result as a `pandas.Series` in GW. **(2 points)**

<!-- grade_id: cell-a4c9cd10fb6f69d5, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: 2-c-t2, points: 2, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
assert isinstance(_, pd.Series)
```

**d)** Plot the two resiual load time series (wind-only and solar-only) for March in a single figure with appropriate labels. **(3 points)**

<!-- grade_id: 2-d-t1, points: 3, grade: true, solution: true -->
```python
# YOUR CODE HERE
```

**e)** Calculate the state of charge time series for (1) the wind-only case and (2) the solar-only case. For each case, choose the initial state of charge to be as small as possible while ensuring that the state of charge does not become negative.

1. Wind-only case. Return the result as a `pandas.Series` in GWh. **(2 points)**

<!-- grade_id: cell-5b5a2d5ce0f58c99, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: 2-e-t1, points: 1, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
assert isinstance(_, pd.Series)
```

<!-- grade_id: 2-e-t2, points: 1, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
```

2. Solar-only case. Return the result as a `pandas.Series` in GWh. **(2 points)**

<!-- grade_id: cell-3ebc50bb90c372ab, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: 2-e-t3, points: 1, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
assert isinstance(_, pd.Series)
```

<!-- grade_id: 2-e-t4, points: 1, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
```

**f)** Plot the state of charge profiles (in TWh) for the wind-only and solar-only cases in a single figure with appropriate labels. **(3 points)**

<!-- grade_id: 2-f-t1, points: 3, grade: true, solution: true -->
```python
# YOUR CODE HERE
```

**g)** Determine the required energy storage capacity (in GWh), the charging power capacity (in GW), and the discharging power capacity (in GW) for (1) the wind-only case and (2) the solar-only case.

1. Wind-only case, energy storage capacity (in GWh). Return the value as a float. **(1 point)**

<!-- grade_id: cell-a60fae33ee3d7c36, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: 2-g-t1, points: 1, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
assert isinstance(_, float)
```

2. Wind-only case, charging power capacity (in GW). Return the value as a positive float. **(1 point)**

<!-- grade_id: cell-cb4420ae527e5156, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: 2-g-t2, points: 1, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
assert isinstance(_, float)
```

3. Wind-only case, discharging power capacity (in GW). Return the value as a positive float. **(1 point)**

<!-- grade_id: cell-4b1ffa5cac9e5144, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: 2-g-t3, points: 1, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
assert isinstance(_, float)
```

4. Solar-only case, energy storage capacity (in GWh). Return the value as a float. **(1 point)**

<!-- grade_id: cell-72cfa2b7d4aa3916, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: 2-g-t4, points: 1, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
assert isinstance(_, float)
```

5. Solar-only case, charging power capacity (in GW). Return the value as a positive float. **(1 point)**

<!-- grade_id: cell-c91d36f1011161a7, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: 2-g-t5, points: 1, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
assert isinstance(_, float)
```

6. Solar-only case, discharging power capacity (in GW). Return the value as a positive float. **(1 point)**

<!-- grade_id: cell-4cc201be3d83989c, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: 2-g-t6, points: 1, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
assert isinstance(_, float)
```

Until now, we only considered the two extreme cases of wind-only and solar-only generation. In the following, we are going to analyze combinations of wind and solar generation.

**h)** Write a Python function that computes the residual load as a function of the solar generation share $\alpha \in [0, 1]$, where $\alpha = 0$ corresponds to the wind-only case and $\alpha = 1$ to the solar-only case. The function should take as input the `pandas.DataFrame` `df` and the solar generation share $\alpha$ and return the residual load time series as a `pandas.Series` in GW. **(3 points)**

<!-- grade_id: 2-h-t1, points: 3, grade: true, solution: true -->
```python
def residual_load(df: pd.DataFrame, alpha: float) -> pd.Series:#
    # YOUR CODE HERE
```

**i)** Apply the function from **h)** for solar generation shares $\alpha$ ranging from 0 to 1 in steps of 0.05 and and store the resulting residual load time series in a new `pandas.DataFrame` where each column denotes a particular value for $\alpha$. Return the resulting `pandas.DataFrame`. **(2 points)**

<!-- grade_id: cell-6d6de49a5ce4df99, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: 2-i-t1, points: 1, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
assert isinstance(_, pd.DataFrame)
```

<!-- grade_id: 2-i-t2, points: 1, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
assert isinstance(_, pd.DataFrame)
```

**j)** Plot the required charging power capacity (in GW) and discharging power capacity (in GW) as a function of the solar generation share $\alpha$. Choose appropriate labels. **(3 points)**

<!-- grade_id: 2-j-t1, points: 3, grade: true, solution: true -->
```python
# YOUR CODE HERE
```

**k)** Plot the required energy storage capacity (in TWh) as a function of the solar generation share $\alpha$. Choose appropriate labels. **(3 points)**

<!-- grade_id: 2-k-t1, points: 3, grade: true, solution: true -->
```python
# YOUR CODE HERE
```

**l)** Determine which of the considered generation mixes (i.e. which value of $\alpha$) results in the smallest required energy storage capacity. Return the value of $\alpha$ as a float. **(2 points)**

<!-- grade_id: cell-c49e1c632fda65fb, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: 2-l-t1, points: 2, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
assert isinstance(_, float)
```

**m)** Plot the state of charge profile for the optimal $\alpha$ from **l)** and compare it to the cases with 100% wind or solar in a single figure with appropriate labels. **(3 points)**

<!-- grade_id: 2-m-t1, points: 3, grade: true, solution: true -->
```python
# YOUR CODE HERE
```