# Instructions

**Please follow these instructions carefully to ensure your submission is graded correctly:**

- Fill in **your answers** in any place that says "YOUR CODE HERE" or "YOUR ANSWER HERE" in the provided `.ipynb` files.

- Your answers to code-based questions must be **generated from code**. Hardcoded answers will result in a deduction of points.  

- Do not change the **autograding cells** labelled with `"AUTOGRADING CELL - DO NOT TOUCH!"`. Modifying or deleting them may lead to incorrect grading results.

- All plots must have appropriate **axis labels** including units where applicable. Failing to do so will result in a partial deduction of points.

- Always proide **units** for quantities where applicable. Failing to do so will result in a partial deduction of points.

- The **code quality** will not be graded. However, concise, well-structured code makes it easier to awared partial points in case of errors.

- The quality of **written answers** and **figures** will be considered during grading.

- Do not use additional packages that are not part of the course **environment**. This would lead to errors during autograding.

- Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel $\rightarrow$ Restart) and then **run all cells** (in the menubar, select Cell $\rightarrow$ Run All).

- Submissions must be your own work, **plagiarism** from the web or peers will be checked for and sanctioned.

- Please upload only the provided `.ipynb` files, filled in with your answers.

- Please **do not** put the `.ipynb` files into archives (e.g., `.zip`, `.tar.gz`); also `.py` files will **not** be accepted. Please **do not change the file names** and do not upload additional files.

- Submission on ISIS *(for TU Berlin students)* requires **two-factor authentication**.

---

## Task 1: Land Eligibility for Wind Turbines (28 points)

Perform an eligibility analysis for the land area available for the development of wind parks in the
Grand Duchy of Luxembourg considering the following criteria:

1. Exclude natural protection areas of the [Natura 2000](https://ec.europa.eu/environment/nature/natura2000/index_en.htm) dataset (`url_natura2000`).
2. Exclude areas within 8 km of airports from the [Natural Earth](https://www.naturalearthdata.com/) dataset (`url_airports`).
3. Exclude areas within 500 m of major roads from the [Natural Earth](https://www.naturalearthdata.com/) dataset (`url_roads`).
4. Exclude the following land cover types from the [CORINE Land Cover](https://land.copernicus.eu/pan-european/corine-land-cover) dataset (`url_corine`, EPSG:3035), including a 900 meter buffer around them:
   - Continuous urban fabric
   - Discontinuous urban fabric
   - Industrial or commercial units
   - Port areas
   - Mineral extraction sites
   - Dump sites
   - Construction sites
   - Green urban areas
   - Sport and leisure facilities
5. Include **only** the following land cover types from the [CORINE Land Cover](https://land.copernicus.eu/pan-european/corine-land-cover) dataset (`url_corine`, EPSG:3035):
   - Non-irrigated arable land
   - Permanently irrigated land
   - Pastures
   - Annual crops associated with permanent crops
   - Land principally occupied by agriculture with significant areas of natural vegetation
   - Natural grasslands
   - Sparsely vegetated areas

*Refer to the course materials to recall how the land cover types are encoded in the CORINE dataset.*

```python
url_natura2000 = "https://tubcloud.tu-berlin.de/s/zPMqJFPD8tKq2Ss/download/Natura2000_end2021-LU.gpkg"
url_airports = (
    "https://tubcloud.tu-berlin.de/s/TbEJ9Lsy9EpTcQS/download/ne_10m_airports.gpkg"
)
url_roads = "https://tubcloud.tu-berlin.de/s/zcKQ95TgpyMJb8E/download/ne_10m_roads.gpkg"
url_corine = "https://tubcloud.tu-berlin.de/s/z7aY8HNCdETQMT7/download/U2018_CLC2018_V2020_20u1-LU.tif"
url_countries = (
    "https://tubcloud.tu-berlin.de/s/7bpHrAkjMT3ADSr/download/country_shapes.geojson"
)
```

```python
import geopandas as gpd
import matplotlib.pyplot as plt
import numpy as np
import pandas as pd
from atlite.gis import ExclusionContainer, shape_availability
from rasterio.plot import show
```

**a)** Load the country shapes as `geopandas.GeoDataFrame`, reduce it to a single entry for Luxembourg and the `geometry` column, and reproject it to the coordinate reference system EPSG:3035. Return the resulting `geopandas.GeoSeries`. **(3 points)**

<!-- grade_id: cell-f8b664556169b0ac, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: 1-a-t1, points: 2, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
assert isinstance(_, gpd.GeoSeries)
```

<!-- grade_id: 1-a-t2, points: 1, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
```


**b)** Calculate the share (in percent) of the area of Luxembourg that is _excluded_ by each of the exclusion criteria 1.-5. individually, using `atlite` functionality. Return the value as a float. **(7 points)**

1. Natura 2000 areas

<!-- grade_id: cell-acc01acc5c0fc87b, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: 1-b-t1, points: 0.5, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
assert isinstance(_, float)
```

<!-- grade_id: 1-b-t2, points: 0.5, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
```

2. Airports

<!-- grade_id: cell-8af46cc4afcd584e, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: 1-b-t3, points: 0.5, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
assert isinstance(_, float)
```

<!-- grade_id: 1-b-t4, points: 0.5, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
```

3. Major roads

<!-- grade_id: cell-8fe62bf1c6843f93, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: 1-b-t5, points: 0.5, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
assert isinstance(_, float)
```

<!-- grade_id: 1-b-t6, points: 0.5, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
```

4. CORINE land classes (exclusion)

<!-- grade_id: cell-22869dd357bc888e, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: 1-b-t7, points: 1, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
assert isinstance(_, float)
```

<!-- grade_id: 1-b-t8, points: 1, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
```

5. CORINE land classes (inclusion)

<!-- grade_id: cell-8c87cfc9058c32af, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: 1-b-t9, points: 1, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
assert isinstance(_, float)

```

<!-- grade_id: 1-b-t10, points: 1, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
```

**c)** Calculate the share (in percent) of the area of Luxembourg that is _eligible_ for wind park development when considering all criteria at once, using `atlite` functionality. Return the value as a float. **(3 points)**

<!-- grade_id: cell-c0864e52a61f9539, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: 1-c-t1, points: 1.5, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
assert isinstance(_, float)
```

<!-- grade_id: 1-c-t2, points: 1.5, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
```

**d)** Visualize the eligible areas on a map of Luxembourg at 100m resolution, using `atlite` functionality. **(3 points)**

<!-- grade_id: 1-d-t1, points: 3, grade: true, solution: true -->
```python
# YOUR CODE HERE
```

**e)** Assume that for the expansion of wind power in Luxembourg, a capacity density of 3 MW/km² of eligible land can be achieved and that the wind conditions result in an average capacity factor of 34%. What is the total electricity (in TWh) that could be produced from wind power in Luxembourg each year? Return the value as a float. **(5 points)**

<!-- grade_id: cell-b4efe2ea2782b02e, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: 1-e-t1, points: 5, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
assert isinstance(_, float)
```

**f)** Perform a sensitivity analysis on the distance criterion to CORINE land classes by repeating the land eligibility calculation with distance requirements of 0m, 200m, 400m, 600m, _together with all other criteria_. Return the results as a `pandas.Series` so that the index contains the distance values (in meters) and the values contain the corresponding share of eligible area (in percent). **(7 points)**

<!-- grade_id: cell-52aee03b9b3b4cc5, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-cc164005ba92a3b3, points: 2.5, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
assert isinstance(_, pd.Series)
```

<!-- grade_id: 1-f-t2, points: 4.5, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
```