# Instructions

**Please follow these instructions carefully to ensure your submission is graded correctly:**

- Fill in **your answers** in any place that says "YOUR CODE HERE" or "YOUR ANSWER HERE" in the provided `.ipynb` files.

- Your answers to code-based questions must be **generated from code**. Hardcoded answers will result in a deduction of points.  

- Do not change the **autograding cells** labelled with `"AUTOGRADING CELL - DO NOT TOUCH!"`. Modifying or deleting them may lead to incorrect grading results.

- All plots must have appropriate **axis labels** including units where applicable. Failing to do so will result in a partial deduction of points.

- Always proide **units** for quantities where applicable. Failing to do so will result in a partial deduction of points.

- The **code quality** will not be graded. However, concise, well-structured code makes it easier to awared partial points in case of errors.

- The quality of **written answers** and **figures** will be considered during grading.

- Do not use additional packages that are not part of the course **environment**. This would lead to errors during autograding.

- Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel $\rightarrow$ Restart) and then **run all cells** (in the menubar, select Cell $\rightarrow$ Run All).

- Submissions must be your own work, **plagiarism** from the web or peers will be checked for and sanctioned.

- Please upload only the provided `.ipynb` files, filled in with your answers.

- Please **do not** put the `.ipynb` files into archives (e.g., `.zip`, `.tar.gz`); also `.py` files will **not** be accepted. Please **do not change the file names** and do not upload additional files.

- Submission on ISIS *(for TU Berlin students)* requires **two-factor authentication**.

---

## Task 2: Merit Order of Electricity Market (33 points)

In this task you are asked to build and evaluate merit order curves for a stylised German day-ahead electricity market in a particular hour.
To solve this task, a dataset on operational power plants and some additional technology-specific data is provided:

```python
url_powerplants = (
    "https://tubcloud.tu-berlin.de/s/icMJ8c89RPpaPFR/download/powerplants.csv"
)
url_technologies = (
    "https://tubcloud.tu-berlin.de/s/XjtnxyNPtPP6eDQ/download/technologies.csv"
)
```

The attributes contained in the two CSV files have the following units:

| Attribute      | Description                          | Unit                           |
|----------------|--------------------------------------|--------------------------------|
| `carrier`        | technology                           |                              |
| `co2_emissions`  | specific carbon dioxide emissions    | t/MWh(thermal)                |
| `color`          | HEX color code                       |                               |
| `efficiency`     | conversion efficiency                | MWh(electric) / MWh(thermal)   |
| `marginal_cost`  | short-term marginal generation cost  | €/MWh(electric)               |
| `p_max_pu`       | capacity factor in particular hour   | p.u.                           |
| `p_nom`          | nominal capacity                     | MW                             |


Assume that all hydro storage has sufficient energy filling levels to dispatch at full capacity.

```python
import matplotlib.pyplot as plt
import numpy as np
import pandas as pd
```

**a)** Read the powerplants dataset (i.e. `powerplants.csv`) into a `pandas.DataFrame` called `df` such that the first column is used as index. Return `df`. **(1 point)**

<!-- grade_id: cell-c76d44520cf544a0, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: 2-a-t1, points: 1, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
assert isinstance(df, pd.DataFrame)
```

**b)** Read the technologies (i.e. `technologies.csv`) as another `pandas.DataFrame` and map the specific carbon dioxide emissions to the powerplants `df` from part **a)** as a new column `co2_emissions` based on the `carrier` attribute. Return the updated `df`.**(2 points)**

<!-- grade_id: cell-8dd0c4fe9b967cb3, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: 2-b-t1, points: 2, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
assert isinstance(df, pd.DataFrame)
```

**c)** For each powerplant, compute the available capacity in MW for the given hour. Store the result in a new column `available_capacity`. Return the updated `df`. **(1 point)**

<!-- grade_id: cell-4e921465f98d8080, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: 2-c-t1, points: 1, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
assert isinstance(df, pd.DataFrame)
```

**d)** Identify the marginal power plant for a total demand of 70 GW in the given hour. Return the index of the marginal power plant as a string. **(4 points)**

<!-- grade_id: cell-a23857fab01cabf8, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: 2-d-t1, points: 4, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
assert isinstance(_, str)
```

**e)** What is the resulting market clearing price in €/MWh? Return the market clearing price as a `float`. **(2 points)** 

<!-- grade_id: cell-a38c79725369e3f7, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: 2-e-t1, points: 2, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
assert isinstance(_, float)
```

**f)** Compute the dispatch (in MW) for each power plant given the total demand of 70 GW in the given hour. Store the result in a new column `dispatch_1`. Return the updated `df`. **(2 points)**

<!-- grade_id: cell-14d144ff126949e6, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: 2-f-t1, points: 1, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
assert isinstance(df, pd.DataFrame)
```

<!-- grade_id: 2-f-t2, points: 1, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
```

**g)** Calculate the total power dispatched per technology. Return a `pandas.Series` with the technology as index and the total dispatched power as values (in MW). **(2 points)**

<!-- grade_id: cell-75cc00e767262ee6, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: 2-g-t1, points: 2, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
assert isinstance(_, pd.Series)
```

**h)** Compute the total profit per technology given the market clearing price. Return a `pandas.Series` with the technology group as index and the total profit as values (in €). **(3 points)**

<!-- grade_id: cell-59313825b2aea140, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: 2-h-t1, points: 3, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
assert isinstance(_, pd.Series)
```

**i)** Calculate the carbon intensity of the system (in gCO2/kWh). Return a single `float` value. **(3 points)**

<!-- grade_id: cell-8399dab3c2ad9aa2, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: 2-i-t1, points: 3, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
assert isinstance(_, float)
```

**In the following questions, consider that a carbon price of 60 €/tCO2 is implemented.**

**j)** Calculate the new short-term marginal generation cost per power plant in €/MWh. Store the result in a new column `marginal_cost_2`. Return the updated `df`. **(3 points)**

<!-- grade_id: cell-89110b77cc3ce22f, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: 2-j-t1, points: 1.5, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
assert isinstance(df, pd.DataFrame)
```

<!-- grade_id: 2-j-t2, points: 1.5, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
```

**k)** What is the new resulting market clearing price in €/MWh? Return the market clearing price as a  `float`. **(1 point)** 

<!-- grade_id: cell-e0c3f729bf26dc03, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: 2-k-t1, points: 1, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
assert isinstance(_, float)
```

**l)** Calculate the absolute difference (in MW) in total power dispatched per technology compared to task **g)**. Negative values denote a decrease in dispatched power, positive values denote an increase. Return a `pandas.Series` with *all* technologies as index. **(2 points)**

<!-- grade_id: cell-0091e4f71e35014c, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: 2-l-t1, points: 2, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
assert isinstance(_, pd.Series)
```

**m)** Calculate the absolute difference (in €) in profits per technology compared to task **h)**. Negative values denote a decrease in profits, positive values denote an increase. Return a `pandas.Series` with *all* technologies as index. **(2 points)**

<!-- grade_id: cell-65f75cf18150ae9a, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: 2-m-t1, points: 2, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
assert isinstance(_, pd.Series)
```

**n)** Calculate the new carbon intensity of the system (in gCO2/kWh). Return a single `float` value. **(1 point)**

<!-- grade_id: cell-c16cc1e48e9bfb8d, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: 2-n-t1, points: 1, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
assert isinstance(_, float)
```

**o)** Plot the two merit order curves (i.e. before and after implementing the carbon price) as step-wise line charts in a single plot. The x-axis should represent the cumulative capacity in MW and the y-axis the marginal cost in €/MWh. The technologies do not need to be labelled. **(4 points)**

<!-- grade_id: 2-o-t1, points: 4, grade: true, solution: true -->
```python
# YOUR CODE HERE
```

```python

```