# Instructions

**Please follow these instructions carefully to ensure your submission is graded correctly:**

- Fill in **your answers** in any place that says "YOUR CODE HERE" or "YOUR ANSWER HERE" in the provided `.ipynb` files.

- Your answers to code-based questions must be **generated from code**. Hardcoded answers will result in a deduction of points.  

- Do not change the **autograding cells** labelled with `"AUTOGRADING CELL - DO NOT TOUCH!"`. Modifying or deleting them may lead to incorrect grading results.

- All plots must have appropriate **axis labels** including units where applicable. Failing to do so will result in a partial deduction of points.

- Always proide **units** for quantities where applicable. Failing to do so will result in a partial deduction of points.

- The **code quality** will not be graded. However, concise, well-structured code makes it easier to awared partial points in case of errors.

- The quality of **written answers** and **figures** will be considered during grading.

- Do not use additional packages that are not part of the course **environment**. This would lead to errors during autograding.

- Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel $\rightarrow$ Restart) and then **run all cells** (in the menubar, select Cell $\rightarrow$ Run All).

- Submissions must be your own work, **plagiarism** from the web or peers will be checked for and sanctioned.

- Please upload only the provided `.ipynb` files, filled in with your answers.

- Please **do not** put the `.ipynb` files into archives (e.g., `.zip`, `.tar.gz`); also `.py` files will **not** be accepted. Please **do not change the file names** and do not upload additional files.

- Submission on ISIS *(for TU Berlin students)* requires **two-factor authentication**.

---

## Task 3: Optimisation for Modelling Electricity Markets (36 points)

Build a simple electricity market model for minimising operational costs within technical constraints for three regions in one particular hour considering the following information:

The operational fleet of power plants in the three regions is specified as follows:

| Name | Technology     | Region     | Marginal Cost [€/MWh] | Capacity [MW] |
| ----- | -------------- | ---------- | --------------------- | ------------- |
| Coal A | Coal           | A | 35                    | 8000         |
| RES A| RES | A | 1                     |  1000         |
| Gas A | Gas            | A | 80                    | 4000         |
| Hydro B | Hydro          | B | 2                     | 4000          |
| Gas B | Gas            | B | 70                    | 2000          |
| RES C | RES            | C | 0                     | 3000          |

The electricity demand in the regions reads as follows:

| Region    | Demand [MW] |
| ---------- | ----------- |
| A | 10000       |
| B | 4000        |
| C | 1000         |

The transmission capacities read as follows:

| Start      | End        | Capacity [MW] |
| ---------- | ---------- | ------------- |
| A | B | 2000           |
| B | C | 300            |
| A | C | 4000           |

Assume equal non-zero reactances for the transmission lines.

**Note: None of the answers can be hard-coded numbers. All answers must be derived from the model results.**

```python
import linopy
import pypsa
import numpy as np
import pandas as pd
```

```python
# marginal costs in EUR/MWh
marginal_costs = pd.DataFrame({
    "A": {"Coal": 35, "RES": 1, "Gas": 80},
    "B": {"Gas": 70, "Hydro": 2},
    "C": {"RES": 0},
})
marginal_costs.columns.name = "region"
marginal_costs.index.name = "technology"

# power plant capacities (nominal powers in MW) in each country
capacities = pd.DataFrame({
    "A": {"Coal": 8000, "RES": 1000, "Gas": 4000},
    "B": {"Gas": 2000, "Hydro": 4000},
    "C": {"RES": 3000},
})
capacities.columns.name = "region"
capacities.index.name = "technology"

# transmission capacities in MW
transmission = pd.Series({"A-B": 2000, "A-C": 4000, "B-C": 300})
transmission.index.name = "transmission"

# country electrical loads in MW
loads = pd.Series({"A": 10000, "B": 4000, "C": 1000}, name="region")
```

**a)** Build and solve the optimisation model using `linopy` by following the steps outlined below.

**a.1)** Create a `linopy.Model` instance under the name `m`. **(1 point)**

**a.2)** Create all variables for generation (`m.g`) and transmission (`m.f`). **(2 points)**

**a.3)** Create the objective function to minimise total operational costs. **(2 points)**

**a.4)** Build the necessary constraints for generation limits, transmission limits and the Kirchhoff laws. **(6 points)**

**a.5)** Solve the model with HiGHS. **(1 point)**

<!-- grade_id: 2-a-t15, points: 12, grade: true, solution: true -->
```python
# YOUR CODE HERE
```

**a.6)** Return the total operational cost as a `float`. **(1 point)**

<!-- grade_id: cell-32d717b768eca961, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: 2-a-t6, points: 1, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
assert isinstance(_, float)
```


**a.7)** Return the optimal generator dispatch as a `pandas.DataFrame` where the index represents the technology of the generator (e.g. `Coal`), the columns represent the region (e.g. `A`), and the values represent the optimal dispatch in MW. **(1 point)**

<!-- grade_id: cell-3dd252ee71cfd437, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: 2-a-at7, points: 1, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
assert isinstance(_, pd.DataFrame)
assert _.columns.symmetric_difference(["A", "B", "C"]).empty
assert _.index.symmetric_difference(["Coal", "Gas", "Hydro", "RES"]).empty
```

**a.8)** Return the optimal power flows as a `pandas.Series` where the index represents the name of the transmission line (e.g. `A-B`) and the values represent the optimal power flow in MW. **(1 point)**

<!-- grade_id: cell-eedcc3e4460a677f, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: 2-a-t8, points: 1, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
assert isinstance(_, pd.Series)
assert _.index.symmetric_difference(transmission.index).empty
```

**a.9)** Return the locational marginal prices (LMPs) as a `pandas.Series` where the index represents the name of the region (e.g. `A`) and the values represent the LMPs in €/MWh. **(1 point)**

<!-- grade_id: cell-1c5cf978a9f16eaa, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: 2-a-t9, points: 1, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
assert isinstance(_, pd.Series)
assert _.index.symmetric_difference(loads.index).empty
```

**b)** Build and solve the optimisation model using `pypsa` by following the steps outlined below.

**b.1)** Create a `pypsa.Network` instance under the name `n`. **(1 point)**

**b.2)** Add the generators, lines, and loads to the network. **(5 points)**

**b.3)** Solve the network with HiGHS. **(1 point)**

<!-- grade_id: 2-b-t13, points: 7, grade: true, solution: true -->
```python
# YOUR CODE HERE
```

**b.4)** Return the total operational cost as a `float`. *(Hint: It should be the same as in the previous task.)* **(1 point)**

<!-- grade_id: cell-805ce0426a6608c7, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: 2-b-t4, points: 1, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
assert isinstance(_, float)
```


**b.5)** Return the optimal generator dispatch as a `pandas.Series` where the index represents the name of the generator (e.g. `Coal A`) and the values represent the optimal dispatch in MW. **(1 point)**

<!-- grade_id: cell-c3adfb061cc6dc11, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: 2-b-t5, points: 1, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
assert isinstance(_, pd.Series)
expected_index = [
    tech + " " + region
    for (tech, region) in capacities.stack().index
]
assert _.index.symmetric_difference(expected_index).empty
```

**b.6)** Return the optimal power flows as a `pandas.Series` where the index represents the name of the transmission line (e.g. `A-B`) and the values represent the optimal power flow in MW. **(1 point)**

<!-- grade_id: cell-eb44f098ce3c495f, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: 2-b-t6, points: 1, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
assert isinstance(_, pd.Series)
assert _.index.symmetric_difference(transmission.index).empty
```

**b.7)** Return the locational marginal prices (LMPs) as a `pandas.Series` where the index represents the name of the region (e.g. `A`) and the values represent the LMPs in €/MWh. **(1 point)**

<!-- grade_id: cell-9434c187add8de10, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: 2-b-t7, points: 1, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
assert isinstance(_, pd.Series)
assert _.index.symmetric_difference(loads.index).empty
```

**c)** Now remove the line between region B and C from the `pypsa` network and re-solve the optimisation problem. **(2 points)**


<!-- grade_id: cell-d6e3b18e0a87e0db, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: 2-c-t0, points: 2, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
```

**c.1)** Return the absolute difference in total operational cost (in €) compared to the solution in **b)** as a `float`. Positive values indicate an increase in cost; negative values indicate a decreased costs. **(1 point)**

<!-- grade_id: cell-b9ee9f2daf3dba9f, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: 2-c-t1, points: 1, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
assert isinstance(_, float)
```

**c.2)** Return the optimal power flows as a `pandas.Series` where the index represents the name of the transmission line (e.g. `A-B`) and the values represent the optimal power flow in MW. **(1 point)**

<!-- grade_id: cell-2121a855aa2e5977, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: 2-c-t2, points: 1, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
assert isinstance(_, pd.Series)
assert _.index.symmetric_difference(["A-B", "A-C"]).empty
```

**c.3)** Return the locational marginal prices (LMPs) as a `pandas.Series` where the index represents the name of the region (e.g. `A`) and the values represent the LMPs in €/MWh. **(1 point)**

<!-- grade_id: cell-e104c70d50f37e57, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: 2-c-t3, points: 1, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
assert isinstance(_, pd.Series)
assert _.index.symmetric_difference(loads.index).empty
```

**c.4)** What is the main reason for the change in operational costs? Describe in 1-2 sentences. **(4 points)**

<!-- grade_id: 2-c-t4, points: 4, grade: true, solution: true -->
YOUR ANSWER HERE

```python

```