# Instructions

**Please follow these instructions carefully to ensure your submission is graded correctly:**

- Fill in **your answers** in any place that says "YOUR CODE HERE" or "YOUR ANSWER HERE" in the provided `.ipynb` files.

- Your answers to code-based questions must be **generated from code**. Hardcoded answers will result in a deduction of points.  

- Do not change the **autograding cells** labelled with `"AUTOGRADING CELL - DO NOT TOUCH!"`. Modifying or deleting them may lead to incorrect grading results.

- All plots must have appropriate **axis labels** including units where applicable. Failing to do so will result in a partial deduction of points.

- Always proide **units** for quantities where applicable. Failing to do so will result in a partial deduction of points.

- The **code quality** will not be graded. However, concise, well-structured code makes it easier to awared partial points in case of errors.

- The quality of **written answers** and **figures** will be considered during grading.

- Do not use additional packages that are not part of the course **environment**. This would lead to errors during autograding.

- Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel $\rightarrow$ Restart) and then **run all cells** (in the menubar, select Cell $\rightarrow$ Run All).

- Submissions must be your own work, **plagiarism** from the web or peers will be checked for and sanctioned.

- Please upload only the provided `.ipynb` files, filled in with your answers.

- Please **do not** put the `.ipynb` files into archives (e.g., `.zip`, `.tar.gz`); also `.py` files will **not** be accepted. Please **do not change the file names** and do not upload additional files.

- Submission on ISIS *(for TU Berlin students)* requires **two-factor authentication**.

---

## Task 1: European Gas Transmission Network (21 points)

The electricity transmission network is not the only network for energy transport in Europe. In this task, you will perform some network analysis on data about the European natural gas transmission network as provided by the [SciGRID_gas](https://web.archive.org/web/20241210201605/http://gas.scigrid.de/) project:

```python
url_nodes = "https://tubcloud.tu-berlin.de/s/8SMwwWQyn6GiPez/download/scigrid-gas-nodes.geojson"

url_pipelines = "https://tubcloud.tu-berlin.de/s/fF6KKpWtJyS3BmD/download/scigrid-gas-pipelines.geojson"
```

```python
import cartopy
import cartopy.crs as ccrs
import geopandas as gpd
import matplotlib.pyplot as plt
import numpy as np
import networkx as nx
import pandas as pd
```

**a)** Read the edges and nodes representing gas transmission pipelines into two separate `geopandas.GeoDataFrame` objects called `nodes` and `edges`. **(2 points)**

<!-- grade_id: cell-30281a37f1f278fd, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: 1-a-t1, points: 2, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
assert isinstance(nodes, gpd.GeoDataFrame)
assert isinstance(edges, gpd.GeoDataFrame)
```

**b)** Plot the transmission network data on a map, such that the colors of the line   indicate the capacity of the pipelines in million cubic metres per day and the width of the lines is proportional to the rated pipeline pressure. The figure should have a size of 18 by 12 inches, a valid projection, and a
colorbar for the colors showing the pipeline capacity. A legend for the line width is not required. Add coastlines and country borders for orientation. **(4 points)**

<!-- grade_id: 1-b-t2, points: 4, grade: true, solution: true -->
```python
# YOUR CODE HERE
```

**c)** Identify one or multiple pipelines with the highest maximum pressure. Return the ID of these pipelines as a list. **(2 points)**

<!-- grade_id: cell-42fe52d7c3d43192, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: 1-c-t1, points: 2, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
assert isinstance(_, list)
```

**d)** What percentage of pipelines crosses borders? Return the answer as a `float` between 0 and 100. **(4 points)**

<!-- grade_id: cell-791638579e5bc7b1, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: 1-d-t1, points: 2, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
assert isinstance(_, float)
```

<!-- grade_id: 1-d-t2, points: 2, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
```

**e)** Create a `networkx.Graph` from the pipeline data using the function `networkx.from_pandas_edgelist()`. As attributes include the length, diameter, capacity and pressure of the pipeline. Return the `networkx.Graph` object. **(3 points)**

<!-- grade_id: cell-6284f9d15a861e91, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: 1-e-t1, points: 3, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
assert isinstance(_, nx.Graph)
```

**f)** Compute the directed incidence matrix of the network and return it as a `numpy.array`. **(2 points)**

<!-- grade_id: cell-300a0e69b9f07bd5, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: 1-f-t1, points: 2, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
assert isinstance(_, np.ndarray)
```

**g)** Provide the frequency distribution of the node degrees in the network as a `pandas.Series` where the index represents the degree and the values represent the number of nodes with that degree. **(2 points)**

<!-- grade_id: cell-c9c07c91ac6db484, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: 1-g-t1, points: 2, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
assert isinstance(_, pd.Series)
```

**h)** Identify the top four nodes with the highest degree. Return their ID as a list. **(2 points)**

<!-- grade_id: cell-31f97f90713b98dc, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: 1-h-t1, points: 2, grade: true, solution: false -->
```python
### AUTOGRADING CELL - DO NOT TOUCH!
assert isinstance(_, list)
```

```python

```